<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/data/Data_Analysis.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


# Kata: Data Analysis

Inspired by an activity in the [Open Machine Learning Course](https://mlcourse.ai/).

## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.

## Part 1 : Loading data

```python
# Import required libraries
import numpy as np
import pandas as pd
```

### Question

Use Pandas to load the dataset called **athlete_events.csv**, which you have to extract from this [Kaggle competition](https://www.kaggle.com/heesoo37/120-years-of-olympic-history-athletes-and-results/). Store it into a DataFrame named `df_olympics`.

<!-- grade_id: cell-6930168ca952d16e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

df_olympics = pd.read_csv("./data/athlete_events.csv")
```

<!-- grade_id: cell-a4b2d9d694a4d11b, points: 1, grade: true, solution: false -->
```python
print(f'df_olympics: {df_olympics.shape}')
assert isinstance(df_olympics, pd.DataFrame)
assert df_olympics.shape == (271116, 15)
```

Output:
```
df_olympics: (271116, 15)
```

## Part 2: Exploring data

### Question

Print general information about the DataFrame.

<!-- grade_id: cell-cc980d2fe4eb912e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

df_olympics.info
```

Output:
```
<bound method DataFrame.info of             ID                      Name Sex   Age  Height  Weight  \
0            1                 A Dijiang   M  24.0   180.0    80.0   
1            2                  A Lamusi   M  23.0   170.0    60.0   
2            3       Gunnar Nielsen Aaby   M  24.0     NaN     NaN   
3            4      Edgar Lindenau Aabye   M  34.0     NaN     NaN   
4            5  Christine Jacoba Aaftink   F  21.0   185.0    82.0   
...        ...                       ...  ..   ...     ...     ...   
271111  135569                Andrzej ya   M  29.0   179.0    89.0   
271112  135570                  Piotr ya   M  27.0   176.0    59.0   
271113  135570                  Piotr ya   M  27.0   176.0    59.0   
271114  135571        Tomasz Ireneusz ya   M  30.0   185.0    96.0   
271115  135571        Tomasz Ireneusz ya   M  34.0   185.0    96.0   

                  Team  NOC        Games  Year  Season            City  \
0                China  CHN  1992 Summer  1992  Summe
```

### Question

Print and observe the first 10 records of the dataset.

<!-- grade_id: cell-c2348e0b371280d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

df_olympics.head(10)
```

Output:
```
ID                      Name Sex   Age  Height  Weight            Team  \
0   1                 A Dijiang   M  24.0   180.0    80.0           China   
1   2                  A Lamusi   M  23.0   170.0    60.0           China   
2   3       Gunnar Nielsen Aaby   M  24.0     NaN     NaN         Denmark   
3   4      Edgar Lindenau Aabye   M  34.0     NaN     NaN  Denmark/Sweden   
4   5  Christine Jacoba Aaftink   F  21.0   185.0    82.0     Netherlands   
5   5  Christine Jacoba Aaftink   F  21.0   185.0    82.0     Netherlands   
6   5  Christine Jacoba Aaftink   F  25.0   185.0    82.0     Netherlands   
7   5  Christine Jacoba Aaftink   F  25.0   185.0    82.0     Netherlands   
8   5  Christine Jacoba Aaftink   F  27.0   185.0    82.0     Netherlands   
9   5  Christine Jacoba Aaftink   F  27.0   185.0    82.0     Netherlands   

   NOC        Games  Year  Season         City          Sport  \
0  CHN  1992 Summer  1992  Summer    Barcelona     Basketball   
1  CHN  2012 Summer  2012
```

### Question

Print and observe 5 random records of the dataset.

<!-- grade_id: cell-ef82a27734491ed4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
df_olympics.sample(5)
```

Output:
```
ID                         Name Sex   Age  Height  Weight  \
99926    50557                   Hu Binyuan   M  38.0   180.0    75.0   
172786   86813                  Helmut Nonn   M  23.0   180.0    79.0   
209889  105389                 Nistor andro   M  21.0   169.0    62.0   
27463    14258  Dianne Marie Bowering-Burge   F  21.0   157.0    50.0   
139242   69939              Set Tommy Limby   M  28.0   178.0    68.0   

             Team  NOC        Games  Year  Season            City  \
99926       China  CHN  2016 Summer  2016  Summer  Rio de Janeiro   
172786    Germany  GER  1956 Summer  1956  Summer       Melbourne   
209889    Romania  ROU  1996 Summer  1996  Summer         Atlanta   
27463   Australia  AUS  1964 Summer  1964  Summer           Tokyo   
139242     Sweden  SWE  1976 Winter  1976  Winter       Innsbruck   

                       Sport                                     Event   Medal  
99926               Shooting                Shooting Men's Double Trap     Na
```

```python

```